# Basic linear algebra in Julia
Author: Andreas Noack Jensen (MIT & JuliaComputing) (https://twitter.com/anoackjensen?lang=en)
(with edits from Jane Herriman)

First let's define a random matrix

```julia
using LinearAlgebra
```

```julia
A = rand(1:4,3,3)
```

Output:
```
3×3 Matrix{Int64}:
 4  4  4
 1  4  4
 1  3  1
```

```julia
B = rand(1:5,3,3)
```

Output:
```
3×3 Matrix{Int64}:
 4  1  4
 1  3  1
 4  1  1
```

```julia
C =[1 2 3; 1 3 2; 5 4 3]
```

Output:
```
3×3 Matrix{Int64}:
 1  2  3
 1  3  2
 5  4  3
```

Define a vector of ones

```julia
x = fill(1.0, (3)) # = fill(1.0, 3)
```

Output:
```
3-element Vector{Float64}:
 1.0
 1.0
 1.0
```

```julia
X = fill(2,(3))
```

Output:
```
3-element Vector{Int64}:
 2
 2
 2
```

Notice that $A$ has type Array{Int64,2} but $x$ has type Array{Float64,1}. Julia defines the aliases Vector{Type}=Array{Type,1} and Matrix{Type}=Array{Type,2}. 

Many of the basic operations are the same as in other languages
#### Multiplication

```julia
 x=["x1";"x2";"x3"]
```

Output:
```
3-element Vector{String}:
 "x1"
 "x2"
 "x3"
```

```julia
@time b = A*x
```

Output:
```
@ ./In[13]:1
 [10] eval(m::Module, e::Any)
    @ Core ./boot.jl:489
```

```julia
?@times
```

Output:
```
No documentation found.

  Binding [36m@times[39m does not exist.
```

#### Transposition
As in other languages `A'` is the conjugate transpose, or adjoint

```julia
A'
```

Output:
```
3×3 adjoint(::Matrix{Int64}) with eltype Int64:
 4  1  1
 4  4  3
 4  4  1
```

and we can get the transpose with

```julia
transpose(A)
```

Output:
```
3×3 transpose(::Matrix{Int64}) with eltype Int64:
 4  1  1
 4  4  3
 4  4  1
```

#### Transposed multiplication
Julia allows us to write this without *

```julia
A'A
```

Output:
```
3×3 Matrix{Int64}:
 18  23  21
 23  41  35
 21  35  33
```

#### Solving linear systems 
The problem $Ax=b$ for ***square*** $A$ is solved by the \ function.

```julia
A
```

Output:
```
3×3 Matrix{Int64}:
 4  4  4
 1  4  4
 1  3  1
```

```julia
b
```

Output:
```
3-element Vector{Float64}:
 12.0
  9.0
  5.0
```

```julia
det(A)
```

Output:
```
-24.0
```

```julia
A\b
```

Output:
```
3-element Vector{Float64}:
 1.0
 1.0
 1.0
```

```julia
B=[1 1 -1;2 -1 3]
```

Output:
```
2×3 Matrix{Int64}:
 1   1  -1
 2  -1   3
```

```julia
C = [1 2]'
```

Output:
```
2×1 adjoint(::Matrix{Int64}) with eltype Int64:
 1
 2
```

```julia
B\C
```

Output:
```
3×1 Matrix{Float64}:
 0.894736842105263
 0.2631578947368419
 0.1578947368421052
```

```julia
 v = rand((1:3),3,3)
```

Output:
```
3×3 Matrix{Int64}:
 2  1  3
 2  3  3
 1  3  2
```

```julia
w = [1 3 5]'
```

Output:
```
3×1 adjoint(::Matrix{Int64}) with eltype Int64:
 1
 3
 5
```

```julia
v\w
```

Output:
```
3×1 Matrix{Float64}:
 -6.000000000000002
  0.9999999999999998
  4.000000000000001
```

`A\b` gives us the *least squares solution* if we have an overdetermined linear system (a "tall" matrix)

```julia
Atall = rand(3, 2)
```

Output:
```
3×2 Matrix{Float64}:
 0.0294459  0.686945
 0.327651   0.413217
 0.174013   0.680338
```

```julia
at = rand(4,4)
```

Output:
```
4×4 Matrix{Float64}:
 0.0553611  0.213165  0.679523  0.870428
 0.0737273  0.514169  0.111927  0.965836
 0.549525   0.250485  0.277576  0.23317
 0.721842   0.234862  0.596575  0.212421
```

```julia

```

```julia
A=rand(1:10,4,4)
```

Output:
```
4×4 Matrix{Int64}:
 5   1   1   9
 1  10   8   1
 4   8  10  10
 7   3   3   1
```

```julia
x=rand(1:10,4)
```

Output:
```
4-element Vector{Int64}:
 1
 3
 1
 8
```

```julia
as = rand((2:5),3)'
```

Output:
```
1×3 adjoint(::Vector{Int64}) with eltype Int64:
 5  5  4
```

```julia
Atall\b
```

Output:
```
2-element Vector{Float64}:
  5.159547425439873
 12.61891545301331
```

```julia
@time b=A*x
```

Output:
```
0.000014 seconds (10 allocations: 528 bytes)
4-element Vector{Int64}:
  81
  47
 118
  27
```

and the *minimum norm least squares solution* if we have a rank-deficient least squares problem

```julia
v = rand(3)
rankef = hcat(v, v)
```

Output:
```
3×2 Matrix{Float64}:
 0.0721794  0.0721794
 0.892831   0.892831
 0.997148   0.997148
```

```julia
be=rand(4)
```

Output:
```
4-element Vector{Float64}:
 0.11739611244184833
 0.43585491562647427
 0.18484533910747603
 0.7252410003542045
```

```julia
rankef\be
```

Output:
```
@ In[73]:1
 [5] eval(m::Module, e::Any)
   @ Core ./boot.jl:489
```

Julia also gives us the minimum norm solution when we have an underdetermined solution (a "short" matrix)

```julia
Ashort = rand(2, 3)
```

Output:
```
2×3 Matrix{Float64}:
 0.301114  0.121313   0.834152
 0.927318  0.0764145  0.684161
```

```julia
bshort = rand(2)
```

Output:
```
2-element Vector{Float64}:
 0.6146739617715813
 0.9004024463562778
```

```julia
Ashort\bshort
```

Output:
```
3-element Vector{Float64}:
 0.5845213668704632
 0.06167412425387078
 0.5169133344150587
```

# The LinearAlgebra library

While much of linear algebra is available in Julia by default (as shown above), there's a standard library named `LinearAlgebra` that brings in many more relevant names and functions. In particular, it provides factorizations and some structured matrix types.  As with all packages, you can bring these additional features into your session with a `using LinearAlgebra`.

### Exercises

#### 10.1 
Take the inner product (or "dot" product) of a vector `v` with itself and assign it to variable `dot_v`.



```julia
v = [1,2,3]
```

Output:
```
3-element Vector{Int64}:
 1
 2
 3
```

```julia
dot_v = v'v
```

Output:
```
14
```

<!-- grade_id: cell-913fef9b0d19cd52, points: 1, grade: true, solution: false -->
```julia
@assert dot_v == 14
```

#### 10.2 
Take the outer product of a vector v with itself and assign it to variable `outer_v`

```julia
@assert dot_v == 14
```

```julia
outer_v = v*v'
```

Output:
```
3×3 Matrix{Int64}:
 1  2  3
 2  4  6
 3  6  9
```

```julia
@assert outer_v == [1 2 3
                    2 4 6
                    3 6 9]
```

#### 10.3 
Use [LinearAlgebra.cross](https://docs.julialang.org/en/v1/stdlib/LinearAlgebra/#LinearAlgebra.cross) to compute the cross product of a vector v with itself and assign it to variable `cross_v`

```julia
cross_v = cross(v,v)
```

Output:
```
3-element Vector{Int64}:
 0
 0
 0
```

<!-- grade_id: cell-e6b6970ffe104df5, points: 1, grade: true, solution: false -->
```julia
@assert cross_v == [0, 0, 0]
```

```julia
a=fill(1,(2,2))
```

Output:
```
2×2 Matrix{Int64}:
 1  1
 1  1
```

```julia
using LinearAlgebra
```

```julia
n = 500 
elapsed_time = 1
for jay in 1:4 
    if jay > 1 
        oldtime = elapsed_time 
    end 
    A = rand(n, n) 
    x = rand(n, 1) 
    elapsed_time  = @elapsed begin
    for rep in 1:100 
        b = A * x 
    end 
    end
    matrixsize = n 
    println("Time: ", elapsed_time, "seconds") 
    if jay > 1 
        ratio = elapsed_time / oldtime 
        println("Ratio: ", ratio) 
    end 
    n = 2 * n 
end


```

Output:
```
Time: 0.007799211seconds
Time: 0.034901762seconds
Ratio: 4.47503753905363
Time: 0.076108426seconds
Ratio: 2.1806470974158843
Time: 0.239697777seconds
Ratio: 3.1494249664288154
```

```julia
B=[14/15 2/15;4/15 22/15;16/15 13/15]
```

Output:
```
3×2 Matrix{Float64}:
 0.933333  0.133333
 0.266667  1.46667
 1.06667   0.866667
```

```julia
U,σ,V = svd(B);
```

```julia
U
```

Output:
```
3×2 Matrix{Float64}:
 -0.333333   0.666667
 -0.666667  -0.666667
 -0.666667   0.333333
```

```julia
σ
```

Output:
```
2-element Vector{Float64}:
 1.9999999999999998
 1.0
```

```julia
V
```

Output:
```
2×2 adjoint(::Matrix{Float64}) with eltype Float64:
 -0.6   0.8
 -0.8  -0.6
```

```julia
sv = [1 2;2 3;3 4]
```

Output:
```
3×2 Matrix{Int64}:
 1  2
 2  3
 3  4
```

```julia
u,σ,v = svd(sv)
```

Output:
```
SVD{Float64, Float64, Matrix{Float64}, Vector{Float64}}
U factor:
3×2 Matrix{Float64}:
 -0.338098   0.847952
 -0.550649   0.173547
 -0.7632    -0.500858
singular values:
2-element Vector{Float64}:
 6.546755636442668
 0.3741532262404972
Vt factor:
2×2 Matrix{Float64}:
 -0.569595  -0.821926
 -0.821926   0.569595
```

```julia
u
```

Output:
```
3×2 Matrix{Float64}:
 -0.338098   0.847952
 -0.550649   0.173547
 -0.7632    -0.500858
```

```julia
σ
```

Output:
```
2-element Vector{Float64}:
 6.546755636442668
 0.3741532262404972
```

```julia
v
```

Output:
```
2×2 adjoint(::Matrix{Float64}) with eltype Float64:
 -0.569595  -0.821926
 -0.821926   0.569595
```

```julia
using Pkg
Pkg.status()
```

Output:
```
[32m[1mStatus[22m[39m `~/.julia/environments/v1.12/Project.toml`
  [90m[7073ff75] [39mIJulia v1.31.1
  [90m[91a5bcdd] [39mPlots v1.41.1
  [90m[c751599d] [39mToeplitzMatrices v0.8.5
  [90m[37e2e46d] [39mLinearAlgebra v1.12.0
  [90m[2f01184e] [39mSparseArrays v1.12.0
```

```julia
names(LinearAlgebra)
```

Output:
```
134-element Vector{Symbol}:
 :/
 :Adjoint
 :BLAS
 :Bidiagonal
 :BunchKaufman
 :Cholesky
 :CholeskyPivoted
 :ColumnNorm
 :Diagonal
 :Eigen
 :Factorization
 :GeneralizedEigen
 :GeneralizedSVD
 ⋮
 :svdvals
 :svdvals!
 :sylvester
 :tr
 :transpose
 :transpose!
 :tril
 :tril!
 :triu
 :triu!
 :×
 :⋅
```

```julia
?names
```

Output:
```
search: [0m[1mn[22m[0m[1ma[22m[0m[1mm[22m[0m[1me[22m[0m[1ms[22m [0m[1mn[22m[0m[1ma[22m[0m[1mm[22m[0m[1me[22mof axes values field[0m[1mn[22m[0m[1ma[22m[0m[1mm[22m[0m[1me[22m[0m[1ms[22m falses

  [36mnames(x::Module; all::Bool=false, imported::Bool=false, usings::Bool=false) [31m->[36m Vector{Symbol}[39m

  Get a vector of the public names of a [36mModule[39m, excluding deprecated names. If
  [36mall[39m is true, then the list also includes non-public names defined in the
  module, deprecated names, and compiler-generated names. If [36mimported[39m is true,
  then names explicitly imported from other modules are also included. If
  [36musings[39m is true, then names explicitly imported via [36musing[39m are also included.
  Names are returned in sorted order.

  As a special case, all names defined in [36mMain[39m are considered "public", since
  it is not idiomatic to explicitly mark names from [36mMain[39m as public.

  [90m
```

```julia
using LinearAlgebra
```

```julia
n = 500
for jay in 1:4
    if jay > 1
        oldtime = time()
    end
    
    M = rand(n, n)  # Create an n x n random matrix
    A = M' * M      # Matrix transpose and multiplication
    
    t = time()      # Capture current time
    
    for rep in 1:5
        R = cholesky(A)  # Perform Cholesky decomposition
    end
    
    matrixsize = n
    elapsed_time = time() - t  # Calculate elapsed time
    
    if jay > 1
        ratio = elapsed_time / oldtime  # Calculate the time ratio
        println("Time ratio: $ratio")
    end
    
    n *= 2  # Double the size of the matrix for the next iteration
end

```

Output:
```
Time ratio: 1.5548996585813235e-11
Time ratio: 1.5570963491882365e-10
Time ratio: 5.69223181499803e-10
```

```julia
using Pkg
Pkg.add("ToeplitzMatrices")
```

Output:
```
[32m[1m    Updating[22m[39m registry at `~/.julia/registries/General.toml`
[32m[1m   Resolving[22m[39m package versions...
[32m[1m   Installed[22m[39m ToeplitzMatrices ─ v0.8.4
[32m[1m   Installed[22m[39m DSP ────────────── v0.7.10
[32m[1m    Updating[22m[39m `~/Documents/julia/Introduction-to-Julia-main/Project.toml`
  [90m[c751599d] [39m[92m+ ToeplitzMatrices v0.8.4[39m
[32m[1m    Updating[22m[39m `~/Documents/julia/Introduction-to-Julia-main/Manifest.toml`
  [90m[621f4979] [39m[92m+ AbstractFFTs v1.5.0[39m
  [90m[187b0558] [39m[92m+ ConstructionBase v1.5.8[39m
  [90m[717857b8] [39m[92m+ DSP v0.7.10[39m
  [90m[7a1cc6ca] [39m[92m+ FFTW v1.8.0[39m
  [90m[1a297f60] [39m[92m+ FillArrays v1.13.0[39m
  [90m[c8e1da08] [39m[92m+ IterTools v1.10.0[39m
  [90m[f27b6e38] [39m[92m+ Polynomials v4.0.11[39m
  [90m[efcf1570] [39m[92m+ Setfield v1.1.1[39m
  [90m[276daf66] [39m[92m+ SpecialFunctions v2.4.0[39m
  [90m[c751599d] [39m[92m
```

```julia
using ToeplitzMatrices

# Example: Create a Toeplitz matrix
c = [1, 2, 3]   # first column
r = [1, 4, 5]   # first row
T = Toeplitz(c, r)

println(T)
```

Output:
```
[1 4 5; 2 1 4; 3 2 1]
```

```julia
T
```

Output:
```
3×3 Toeplitz{Int64, Vector{Int64}, Vector{Int64}}:
 1  4  5
 2  1  4
 3  2  1
```

```julia
?Toeplitz
```

Output:
```
search: [0m[1mT[22m[0m[1mo[22m[0m[1me[22m[0m[1mp[22m[0m[1ml[22m[0m[1mi[22m[0m[1mt[22m[0m[1mz[22m [0m[1mT[22m[0m[1mo[22m[0m[1me[22m[0m[1mp[22m[0m[1ml[22m[0m[1mi[22m[0m[1mt[22m[0m[1mz[22mMatrices Symme[0m[1mt[22mricT[0m[1mo[22m[0m[1me[22m[0m[1mp[22m[0m[1ml[22m[0m[1mi[22m[0m[1mt[22m[0m[1mz[22m Abs[0m[1mt[22mractT[0m[1mo[22m[0m[1me[22m[0m[1mp[22m[0m[1ml[22m[0m[1mi[22m[0m[1mt[22m[0m[1mz[22m

[36m  Toeplitz[39m

  A Toeplitz matrix.

  ────────────────────────────────────────────────────────────────────────────

[36m  Toeplitz(vc::AbstractVector, vr::AbstractVector)[39m

  Create a [36mToeplitz[39m matrix from its first column [36mvc[39m and first row [36mvr[39m where
  [36mvc[1] == vr[1][39m.

  ────────────────────────────────────────────────────────────────────────────

[36m  Toeplitz(A::AbstractMatrix)[39m

  "Project" matrix [36mA[39m onto its Toeplitz part using the first row
```

```julia
C = [10 4 3; 4 1 2; 3 2 5]
```

Output:
```
3×3 Matrix{Int64}:
 10  4  3
  4  1  2
  3  2  5
```

```julia
isposdef(C)
```

Output:
```
false
```

```julia
cholesky(C)
```

Output:
```
@ In[112]:1
 [11] eval(m::Module, e::Any)
    @ Core ./boot.jl:489
```

```julia

```

```julia

```